```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

---
## Part 1: Using Jupyter Notebook

### Running Cells and Displaying Output

Run the following cell.  If you are unfamiliar with Jupyter Notebooks, skim [this tutorial](http://nbviewer.jupyter.org/github/jupyter/notebook/blob/master/docs/source/examples/Notebook/Notebook%20Basics.ipynb) or selecting **Help --> JupyterLab Reference** in the menu bar above. 

```python
print("Hello World!")
```

In Jupyter notebooks, all print statements are displayed below the cell. Furthermore, the output of the last line is displayed following the cell upon execution.

```python
"Will this line be displayed?"

print("Hello" + ",", "world!")

5 + 3
```

### Viewing Documentation

To output the documentation for a function, use the `help` function.

```python
help(print)
```

You can also use Jupyter to view function documentation inside your notebook. The function must already be defined in the kernel for this to work.

Below, click your mouse anywhere on `print()` and use `Shift` + `Tab` to view the function's documentation. 

```python
print('Welcome to PSTAT 234.')
```

### Importing Libraries and Magic Commands

In PSTAT 234, we will be using common Python libraries to help us process data. By convention, we import all libraries at the very top of the notebook. There are also a set of standard aliases that are used to shorten the library names. Below are some of the libraries that you may encounter throughout the course, along with their respective aliases.

```python
import pandas as pd
import numpy as np
```

Another useful magic command is `%%time`, which times the execution of that cell. You can use this by writing it as the first line of a cell. (Note that `%%` is used for *cell magic commands* that apply to the entire cell, whereas `%` is used for *line magic commands* that only apply to a single line.)

```python
%%time

lst = []
for i in range(100):
    lst.append(i)
```

### Keyboard Shortcuts

Even if you are familiar with Jupyter, we strongly encourage you to become proficient with keyboard shortcuts (this will save you time in the future). To learn about keyboard shortcuts, go to **Help --> Keyboard Shortcuts** in the menu above. 

Here are a few that we like:
1. `Ctrl` + `Return` : *Evaluate the current cell*
1. `Shift` + `Return`: *Evaluate the current cell and move to the next*
1. `ESC` : *command mode* (may need to press before using any of the commands below)
1. `a` : *create a cell above*
1. `b` : *create a cell below*
1. `dd` : *delete a cell*
1. `z` : *undo the last cell operation*
1. `m` : *convert a cell to markdown*
1. `y` : *convert a cell to code*

### Jupyter Lab

JupyterLab is the next generation Jupyter environment that includes more than Jupyter Notebook.
![JupyterLab](images/jupyterlab.png)

To use Jupyter Lab, visit http://pstat-134-234.lsit.ucsb.edu.

---
# Part 2: Jupyter notebook labs and assignments

**Note: this lab notebook will not count toward your grade, but it is designed to give you a full understanding of how the graders work for this class**

In this class, some labs and assignments based on Jupyter notebooks will be graded by Gradescope autograder. The grading is done by "tests" some of which are visible and others are hidden.

The remainder of this notebook illustrates how the grading works, and, during the process, explain various components that help you understand and work on your labs and assignments easier.

#### Instructions

The software being used is called [**otter grader**](https://otter-grader.readthedocs.io). Here are the basics:

* **You must NOT change the notebook file name.**  
    Notice the first cell has the name of this notebook `lab1.ipynb`. Changing the name of this Jupyter notebook file will cause problems.

* **Each question has a response cell _directly_ below it.**  
    This is where you answer the question. If you have extra content in this cell, that's fine.
    
* **If a question has tests, grading cell follows the response cell.**  
    Most response cells are followed by a test cell that runs automated tests to check your work: e.g.  
    ```
    grader.check("q1");
    ```
    Test results are meant to give you some useful feedback, but it's your responsibility to answer the question. 
    
* **Keep Question-Response-Grading cell intact**  
    Please don't delete questions, response cells, or test cells. You won't get credit for your work if you do.

* **Important: test cells don't always confirm that your response is correct.**  
    There may be other tests that we run when scoring your notebooks. We **strongly recommend** that you check your solutions yourself rather than just relying on the test cells.
    
* **Double check your submission**  
    Make sure the correct file is submitted
    
Please refer to this documentation for how you can use it more effectively: https://otter-grader.readthedocs.io/en/latest/otter_check.html

## Sample question

Suppose the question to be graded is 
> What is the smallest prime number? Assign it to variable `x`.

The correct answer would be `x = 2`.

## Tests

A test in our Jupyter notebook grading context is a piece of code and the expected result of that code. For the sample question, some examples of tests could be, 

1. Test code: `isinstance(x, int)`  
    Expected result: `True`
2. Test code: `0 < x < 100`  
    Expected result: `True`
3. Test code: `print(x)`  
    Expected result: `2`
    
In your notebook, tests such as these are executed to evaluate your grade. Below, the outputs in your notebook and Gradescope are explained in detail.

## Example Question 1

What is the smallest prime number? Assign it to variable `x`.

The correct answer is, 
```python
x = 2
```

Suppose following are the three tests

| Test code  | Expected result |
| -------    | -------         |
| `isinstance(x, int)` | `True`          |
| `0 < x < 100` | `True`          |
| `print(x)` | `2`          |

If the student responds correctly as below, all tests would pass.

```python
x = 2 # student response
```

```python
grader.check("q1")
```

## Example Question 2

We will create a similar question, but student will respond differently to this question.

What is the smallest prime number? Assign it to variable `y`.

The correct answer is, 
```python
y = 2
```

Suppose following are the three tests

| Test code  | Expected result |
| -------    | -------         |
| `isinstance(y, int)` | `True`          |
| `0 < y < 100` | `True`          |
| `print(y)` | `2`          |

If the response is something like `y=2.3`, some tests can pass, although the response is technically wrong.

```python
y = 2.3 # student response
```

```python
grader.check("q2")
```

## Example Question 3

We will create a yet another similar question, but one of the tests will be hidden from you.

What is the smallest prime number? Assign it to variable `z`.

The correct answer is, 
```python
z = 2
```

Suppose following are the three tests

| Test code  | Expected result |
| -------    | -------         |
| `isinstance(z, int)` | `True`          |
| `0 < z < 100` | `True`          |
| `print(z)` | `2`          |

In particular, the third test is hidden from you. If the response is `z=3`, the first two tests would be shown as passing; however, the result of the third test will not be visible when you are working on the notebook.

```python
z = 3 # student response
```

```python
grader.check("q3")
```

## More about tests and expected output

Except for any hidden tests, each test and the corresponding expected output are a part of the notebook. You can inspect them by opening any notebook as a json: right click on `lab1.ipynb` > "Open With" > choose "JSON". Then navigate to `root` > `metadata` > `otter` > `tests` > `q3`, for example.

You can infer that the total number of points for this question is 3, and there are two visible tests. For each test, the test code and expected result are shown as the dictionary value for key, `'code'`. For instance, the two tests/expected results are,

1. `isinstance(z, int)\nTrue`  
    test code: `isinstance(z, int)`  
    expected result: `True`
2. `0 < z < 100\nTrue`
    test code: `0 < z < 100`  
    expected result: `True`

**A word of caution**: the tests being correct is a necessary condition but not sufficient condition for the correctness of your code. In other words, you may be passing all the (visible) tests; however, your response could still be incorrect. In [Example Question 3](#example-question-3), the test of  has a hidden test that checks if the value of `print(z)` is `2`.  You will not see any hidden tests while you are working on the assignment; however, the result of the hidden tests are revealed when final grades are released on Gradescope.

<!-- BEGIN QUESTION -->

## Example Question 4

Here is a question to be graded manually: create a matplotlib visualization of a scatter plot.

```python
%matplotlib inline

import matplotlib.pyplot as plt
import numpy as np

data = {'a': np.arange(50),
        'c': np.random.randint(0, 50, 50),
        'd': np.random.randn(50)}
data['b'] = data['a'] + 10 * np.random.randn(50)
data['d'] = np.abs(data['d']) * 100

plt.scatter('a', 'b', c='c', s='d', data=data)
plt.xlabel('entry a')
plt.ylabel('entry b')
plt.show()
```

<!-- END QUESTION -->

# Part3 Some notes on Text part
## Text

There are some descriptive questions in labs and homework, you need to change the cell to `Markdown` which is the proper enviroment for text. If you type your descriptive answers in code cells directly, the cell can not run. Try to uncomment the following cell and see what happened.

```python
#Test text in a code cell
```

If you type your answer as comments in a code cell, the long sentence will be cut of automatically when you export your submission and it may cause some points deductions.

## LaTeX
[LaTeX](https://www.latex-project.org) is a programming enviroment for producing scientific documents. You will need to use it in homework and projects in this course.

### Mathematics Inline and Display

Enclose LaTeX code in dollar signs `$...$` to display math inline. For example, the code `$\sum_{i = 1}^n (\hat y_i - y_i)^2$` renders inline as $\sum_{i = 1}^n (\hat y_i - y_i)^2$.

Enclose LaTeX code in double dollar signs `$$...$$` to display expressions in a centered paragraph. For example:

`$$f'(a) = \lim_{x \to a} \frac{f(x) - f(a)}{x-a}$$`
renders as
$$f'(a) = \lim_{x \to a} \frac{f(x) - f(a)}{x-a}$$

See the [LaTeX WikiBook](https://en.wikibooks.org/wiki/LaTeX) for more information(especially the mathematics section).

### Common Symbols
Below we give a list of commonly used mathematical symbols.

For any other symbols you can think of, they can be found by [Detexify](http://detexify.kirelabs.org/classify.html).

| Syntax  | Output |
| -------    | -------     |
| `x_n` | $x_n$          |
| `x^2` | $x^2$          |
| `\infty` | $\infty$    |
| `\frac{a}{b}`|$\frac{a}{b}$|
| `\partial` | $\partial$ |
| `\alpha` | $\alpha$ |
| `\sum_{n=0}^{\infty}` | $\sum_{n=0}^{\infty}$|
| `\prod_{n=0}^{\infty}` | $\prod_{n=0}^{\infty}$|
| `\int_a^b` | $\int_a^b$|
| `\mathbb{Z}` | $\mathbb{Z}$|
|`\dots`| $\dots$ |
|`\vdots` | $\vdots$|
|`\ddots`| $\ddots$|
| -------    | -------     |

### Matrices and Brackets

Create a matrix without brackets:

`$$\begin{matrix} a & b \\ c & d \end{matrix}$$`
$$\begin{matrix} a & b \\ c & d \end{matrix}$$

Create a matrix with round brackets:

`$$\begin{pmatrix} a & b \\ c & d \end{pmatrix}$$`
$$\begin{pmatrix} a & b \\ c & d \end{pmatrix}$$

Create a matrix with square brackets:

`$$\begin{bmatrix} 1 & 2 & 1 \\ 3 & 0 & 1 \\ 0 & 2 & 4 \end{bmatrix}$$`
$$\begin{bmatrix} 1 & 2 & 1 \\ 3 & 0 & 1 \\ 0 & 2 & 4 \end{bmatrix}$$

Use `\left` and `\right` to enclose any expression in brackets:

`$$\left( \frac{p}{q} \right)$$`
$$\left( \frac{p}{q} \right)$$

### Example: Jacobian Matrix

The Jacobian matrix of the function $f(x_1, \ldots, x_n)$ is:

`$$
\mathbf{J}
=
\frac{d \mathbf{f}}{d \mathbf{x}}
=
\left[ \frac{\partial \mathbf{f}}{\partial x_1}
\cdots \frac{\partial \mathbf{f}}{\partial x_n} \right]
=
\begin{bmatrix}
\frac{\partial f_1}{\partial x_1} & \cdots &
\frac{\partial f_1}{\partial x_n} \\
\vdots & \ddots & \vdots \\
\frac{\partial f_m}{\partial x_1} & \cdots &
\frac{\partial f_m}{\partial x_n}
\end{bmatrix}
$$`
$$
\mathbf{J}
=
\frac{d \mathbf{f}}{d \mathbf{x}}
=
\left[ \frac{\partial \mathbf{f}}{\partial x_1}
\cdots \frac{\partial \mathbf{f}}{\partial x_n} \right]
=
\begin{bmatrix}
\frac{\partial f_1}{\partial x_1} & \cdots &
\frac{\partial f_1}{\partial x_n} \\
\vdots & \ddots & \vdots \\
\frac{\partial f_m}{\partial x_1} & \cdots &
\frac{\partial f_m}{\partial x_n}
\end{bmatrix}
$$

### Displaying equations on separate lines

The `align` environment in Latex specifies that you want centered equations, separated from the text. It is called `align` because it allows you to align the equations. You separate lines in the equations with a double backslash (`//`). Insert an ampersand (`&`) in each line at the alignment point. All equations will be aligned at the location of the ampersand symbols (and, of course, the ampersands will not appear in the rendered equations).

For a three-vector consisting of $x$, $y$, and $z$ components,

`\begin{align}
\mathbf{a} \cdot \mathbf{b} &= \sum_{i=1}^n a_i b_i \\
&= a_x b_x + a_y b_y + a_z b_z.
\end{align}`

\begin{align}
\mathbf{a} \cdot \mathbf{b} &= \sum_{i=1}^n a_i b_i \\
&= a_x b_x + a_y b_y + a_z b_z.
\end{align}

We can use `array` enviroment to represent the functions in multiple cases.

`$$
P(\sigma) = \left\{
\begin{array}{cl}
\sigma^{-1} & \sigma > 0 \\
0 & \sigma \le 0.
\end{array}
\right.
$$`

$$
P(\sigma) = \left\{
\begin{array}{cl}
\sigma^{-1} & \sigma > 0 \\
0 & \sigma \le 0.
\end{array}
\right.
$$
Another way to represent functions with multiple cases is using `cases` enviroment:
`$$
P(\sigma)=
\begin{cases}
\sigma^{-1} & \sigma > 0 \\
0 & \sigma \le 0.
\end{cases}
$$`
$$
P(\sigma)=
\begin{cases}
\sigma^{-1} & \sigma > 0 \\
0 & \sigma \le 0.
\end{cases}
$$

### Write text in formula

Sometimes, we need to add some explainations in formula, then we can use `\text{}` to insert text. For example,
`\begin{align}
\text{good: } &\int_0^{2\pi} \mathrm{d}x \, \sin x. \\
\text{bad: } &\int_0^{2\pi} \mathrm{d}x \sin x.
\end{align}`
\begin{align}
\text{good: } &\int_0^{2\pi} \mathrm{d}x \, \sin x. \\
\text{bad: } &\int_0^{2\pi} \mathrm{d}x \sin x.
\end{align}

### Using LaTeX in figures

Fortunately, Matplotlib allows use of Latex in axis labels and other text in figures. You simply need to include the dollar signs in your string.

```python
x = np.linspace(0.0, 2.0 * np.pi, 200)
y = np.exp(np.sin(x))
plt.plot(x, y)
plt.xlabel('$x$')
plt.ylabel('$\mathrm{e}^{\sin x}$')
```

## Preparing Gradescope Submission Package

The function `grader.export()` compiles a pdf file with your responses to the manually graded questions. When you are ready to create your submission package, follow these steps:

1. Go to Kernel > Restart Kernel and Run All Cells...
2. Save your notebook
3. Again! Go to Kernel > Restart Kernel and Run All Cells...

You must open the pdf and make sure all your responses show up correctly.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the zip file and submit to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 3)
```
>>> isinstance(x, int)
True
```
```
>>> 0 < x < 100
True
```
```
>>> print(x)
2
```


Tests for q2 (points: 3)
```
>>> isinstance(y, int)
True
```
```
>>> 0 < y < 100
True
```
```
>>> print(y)
2
```


Tests for q3 (points: 3)
```
>>> isinstance(z, int)
True
```
```
>>> 0 < z < 100
True
```
